---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`
121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212

`@hint`
1212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212

`@pre_exercise_code`
```{r}
12121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212
```

`@sample_code`
```{r}

```

`@solution`
```{r}
121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212121212
```

`@sct`
```{r}

```
